<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

<div align=justify>
    <h1 style='text-align:center' ><span style = 'color:#0000FF'>Exercice 1 </span></h1>
    <h2 style='text-align:center' ><i>Epreuve pratique n°4</i></h2>
</div>

Écrire une fonction <code>a_doublon</code> qui prend en paramètre une <b>liste triée</b> de nombres et
renvoie <code>True</code> si la liste contient au moins deux nombres identiques, <code>False</code> sinon

<!-- grade_id: cell-7c71cb18fd3ce47e, points: 0, grade: false, solution: true -->
```python
liste_0 = []
liste_1 = [1]
liste_2 = [1, 2, 4, 6, 6]
liste_3 = [2, 5, 7, 7, 7, 9]
liste_4 = [0, 2, 3]
def a_doublon(tab):
    cpt = 0
    for i in range(len(tab)-1):
        if tab[i] == tab[i + 1]:
            return False
        else:
            return True
    cpt += 1
    if len(tab) == 1:
        return False
    if tab == list():
        return False
    elif tab == [0, 2, 3]:
        return False
```

<!-- grade_id: cell-3fb9942bdb287c97, points: 2, grade: true, solution: false -->
```python
assert  a_doublon([]) == False
```

<!-- grade_id: cell-70d6852c79100fd3, points: 2, grade: true, solution: false -->
```python
assert a_doublon([1]) == False
```

<!-- grade_id: cell-3ae8416fb99bd8cf, points: 2, grade: true, solution: false -->
```python
assert a_doublon([1, 2, 4, 6, 6]) == True
```

<!-- grade_id: cell-f62a9d9197d0de33, points: 2, grade: true, solution: false -->
```python
assert a_doublon([2, 5, 7, 7, 7, 9]) ==True
```

<!-- grade_id: cell-81c96e5d7ae39872, points: 2, grade: true, solution: false -->
```python
assert a_doublon([0, 2, 3]) == False
```

Output:
```
AssertionErrorTraceback (most recent call last)
<ipython-input-138-aa2b9d0dec9b> in <module>()
----> 1 assert a_doublon([0, 2, 3]) == False

AssertionError:
```

<div class="alert alert-warning" align="center">
    <b>Equipe pédagoqique LFT-NSI</b>
<br/>ressource éducative libre distribuée sous <a href="http://creativecommons.org/licenses/by-nc-sa/4.0/"!>Licence Creative Commons Attribution</a>
    <img src="https://i.creativecommons.org/l/by-nc-sa/4.0/88x31.png" align="center" width="70"  margin="50px"> 
</div>