<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

<div align=justify>
    <h1 style='text-align:center' ><span style = 'color:#0000FF'>Exercice 2 </span></h1>
    <h2 style='text-align:center' ><i>Epreuve pratique n°4</i></h2>
</div>

On considére une relation d'ordre dans un tableau <code>ordre</code> de $n$ cases
d’entiers distincts deux à deux et compris entre $1$ et $n$.

Par exemple, <code>ordre = [5, 4, 3, 6, 7, 2, 1, 8, 9]</code> dans le cas $n=9$.

On dit qu’il y a un point de rupture dans <code>ordre</code>  dans chacune des situations suivantes :
<ul>
    <li>la première valeur de ordre n’est pas 1 ;</li>
    <li>l’écart entre deux entiers consécutifs n’est pas égal à 1 ;</li>
<li>la dernière valeur de ordre n’est pas n.</li>
</ul>

Par exemple, si <code>ordre = [5, 4, 3, 6, 7, 2, 1, 8, 9]</code> avec $n = 9$, on a
<ul>
    <li>un point de rupture au début car 5 est différent de 1</li>
    <li>un point de rupture entre 3 et 6 (l’écart est de 3)</li>
    <li>un point de rupture entre 7 et 2 (l’écart est de 5)</li>
    <li>un point de rupture entre 1 et 8 (l’écart est de 7)</li>
</ul>
Il y a donc 4 points de rupture.


Compléter la fonction Python <code>est_un_ordre</code> :

la fonction <code>est_un_ordre</code> renvoie <code>True</code> si le tableau passé en paramètre représente bien un ordre d'entiers et <code>False</code> sinon

```python 
def est_un_ordre(tab):
     '''
     Renvoie True si tab est de longueur n et contient tous les
    entiers de 1 à n, False sinon
     '''
     for i in range(1,...):
         if ...:
             return False
     return True
```

<!-- grade_id: cell-9168b948b50af05d, points: 0, grade: false, solution: true -->
```python
liste_1 = [1, 6, 2, 8, 3, 7]
liste_2 = [5, 4, 3, 6, 7, 2, 1, 8, 9]
def est_un_ordre(tab):
    for i in range(1, len(tab)):
        if tab[0] == 1:
            return False
    return True
```

<!-- grade_id: cell-e05fd7e368fce26f, points: 2, grade: true, solution: false -->
```python
assert est_un_ordre([1, 6, 2, 8, 3, 7]) == False
```

<!-- grade_id: cell-8a4b6c5a7896a11c, points: 2, grade: true, solution: false -->
```python
assert  est_un_ordre([5, 4, 3, 6, 7, 2, 1, 8, 9]) == True
```

Compléter la fonctions Python <code>nombre_points_rupture</code> proposée pour que :
<ul>
    <li>la fonction <code>nombre_points_rupture</code> renvoie le nombre de points de rupture
d’un tableau passé en paramètre.
    </li>
</ul>

```python
def nombre_points_rupture(tab):
    '''
        Renvoie le nombre de point de rupture de tab
        Precondition : tab est un ordre valide 
    '''
    n = len(tab)
    nb = 0
    if tab[...] != 1: 
        nb = nb + 1
    i = 0
    while i < ....:
        if ........ not in [-1, 1]:
            nb = nb + 1
        i = i + 1
    if tab[...] != n:
        nb = nb + 1
    return nb
```

<!-- grade_id: cell-782b5f30cc8b5832, points: 0, grade: false, solution: true -->
```python
def nombre_points_rupture(tab):
    '''
        Renvoie le nombre de point de rupture de tab
        Precondition : tab est un ordre valide 
    '''
    n = len(tab)
    nb = -1 #Dans la consigne nb = 0 mais pas -1
    if  tab[0] != 1:
        nb = nb + 1
    i = 0
    while i < n:
        if tab[i] - tab[i - 1]not in [-1, 1]:
            nb = nb + 1
        i = i + 1
    if tab[n - 1] != n:
        nb = nb + 1
    return nb
```

<!-- grade_id: cell-72d87b3f8733058d, points: 2, grade: true, solution: false -->
```python
assert  nombre_points_rupture([5, 4, 3, 6, 7, 2, 1, 8, 9]) == 4
```

<!-- grade_id: cell-695434cfaba3b9ff, points: 2, grade: true, solution: false -->
```python
assert nombre_points_rupture([1, 2, 3, 4, 5]) == 0
```

<!-- grade_id: cell-cbb0c280ca2d0eda, points: 2, grade: true, solution: false -->
```python
assert nombre_points_rupture([1, 6, 2, 8, 3, 7, 4, 5]) == 7
```

<div class="alert alert-warning" align="center">
    <b>Equipe pédagoqique LFT-NSI</b>
<br/>ressource éducative libre distribuée sous <a href="http://creativecommons.org/licenses/by-nc-sa/4.0/"!>Licence Creative Commons Attribution</a>
    <img src="https://i.creativecommons.org/l/by-nc-sa/4.0/88x31.png" align="center" width="70"  margin="50px"> 
</div>